---
title: "Using Filter"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
    fig_caption: false
    css: css/textbook.min.css
runtime: shiny_prerendered
description: >
  Learn how to filter the rows in a tibble with `filter()`. You will also learn how to chain operations together with the magrittr pipe operator `%>%`.
---

<!-- Known Issues 

  1. The TOC breaks if SASS is compiled in the rmarkdown file
  2. Currently the TOC progress bar isn't progressing correctly - cause unknown, might be css modifications added or shinyjs?
-->

<!-- Code for including shinyJS, currently used for toggling css classes toggleClass() | https://deanattali.com/shinyjs/advanced#usage-prerendered -->
```{r, echo=FALSE}
shiny::addResourcePath("shinyjs", system.file("srcjs", package = "shinyjs"))
```
```{r, context="server"}
shinyjs::useShinyjs(html = TRUE)
```
<script src="shinyjs/inject.js"></script>

<!-- End Code for shinyJs -->
```{r setup, include=FALSE}
library(learnr)
gradethis::gradethis_setup()

##Setup Homework Grade Tracker
hmw_info <- reactiveValues(s_id = "Not Set ", 
                                hmw_grade = 0.0, 
                                #Eventually update this section, so that it can be used to generate the homework questions later on. Without generation,
                                #this section could become out of date/sync with actual problems listed below.
                                #Idea - Have each question append its information to hmw_info$submissions! This would remove the syncing issue.
                                submissions = list(hmw1 = list(type="questionaire", 
                                                               q_text="(True or False), filter() modifies the dataframe passed to it.",
                                                               choices=c("True", "False"), 
                                                               is_correct = FALSE, 
                                                               submitted_answer = ""),
                                                   hmw2 = list(type="questionaire", 
                                                               q_text="How many rows will this code segment return? $\\texttt{filter(mpg, displ > 2)}$",
                                                               choices=c("191", "212", "11", "0"), 
                                                               is_correct = FALSE, 
                                                               submitted_answer = ""),
                                                   hmw3 = list(type="code", 
                                                               q_text="Fill in the blanks with the necessary code so that only the audi cars with a city mpg of 20 or more are listed.", 
                                                               choices="filter(mpg, manufacturer ____ \"audi\" ____ cty ____ 20)",
                                                               is_correct = FALSE, 
                                                               submitted_answer = ""),
                                                   hmw4 = list(type="code", 
                                                               q_text="Debugging Exercise - Fix the following piece of code so that only compact cars are listed in the results.", 
                                                               choices="filter(mpg, CLASS == \"COMPACT\")", 
                                                               is_correct = FALSE, 
                                                               submitted_answer = ""),
                                                   hmw5 = list(type="code", 
                                                               q_text="Write the necessary code to list all of the cars with an average highway mpg of 20 or more.", 
                                                               choices=NULL, 
                                                               is_correct = FALSE, 
                                                               submitted_answer = ""),
                                                   hmw6 = list(type="code", 
                                                               q_text="Write the necessary code to list out all of the minivans and subcompacts produced in 1999.", 
                                                               choices=NULL, 
                                                               is_correct = FALSE, 
                                                               submitted_answer = "")))

new_recorder <- function(tutorial_id, tutorial_version, user_id, event, data) {
   #cat("Original: \n", tutorial_id, " (", tutorial_version, "): ", user_id, ", ", event, "\n", sep = "")
  
  hmw_info$s_id <<- user_id #Set the student's name to match their user id (should be cu id)
  if(!is.null(data$label)){
    if(grepl("hmw", data$label, fixed =TRUE)){ ##Only verify homework problems  
      if(event == "exercise_result"){#ensure we are checking a coding problem after it's been evaluated
          #Verify Coding Problem
          if(!is.null(data$feedback)){#If feedback is present, then this is a submitted problem
            if(data$feedback$correct){
              #Only update for a correct answer once. Ignore repeated correct answers.
              if(!hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$is_correct){
                hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$is_correct <<- TRUE
                hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$submitted_answer <<- data$code
                hmw_info$hmw_grade <<- hmw_info$hmw_grade + 1
              }
            }
            #Check if they haven't answered the question correctly. If not, we want to store their most recent incorrect attempt
            else if(!hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$is_correct){
              hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$submitted_answer <<- data$code
            }
          }
        }
      else if(event == "question_submission"){
        #Check if they answered the interactive question correctly
        if(data$correct){
          #Only update for a correct answer once. Ignore repeated correct answers.
          if(!hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$is_correct){
            hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$is_correct <<- TRUE
            hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$submitted_answer <<- data$answer
            hmw_info$hmw_grade <<- hmw_info$hmw_grade + 1
          }
        }
        #Check if they haven't answered the question correctly. If not, we want to store their most recent incorrect attempt
        else if(!hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$is_correct){
          hmw_info$submissions[[strsplit(data$label, split="_")[[1]][2]]]$submitted_answer <<- data$answer
        }
      } 
    }
  }
  
}

options(tutorial.event_recorder = new_recorder)

library(tidyverse)
library(palmerpenguins)
library(reactable)
#library(sass)
library(shinyjs)
#Compile SASS code into a local css copy
# Warning - Don't include in final version! This breaks/removes the TOC when Shiny-Server serves this Learnr tutorial
# Only run in console  to generate a new css file
#sass(sass_file("../SharedResources/textbook.sass"),
#  output = "css/textbook.min.css",
#  options = sass_options(
#    output_style = "compressed",
#    source_map_embed = TRUE
#  ))



knitr::opts_chunk$set(error = FALSE, message = FALSE)
```


## What is filter()?

### <span style="opacity:0;">filter()</span>

![Cartoon showing three fuzzy monsters either selecting or crossing out rows of a data table. If the type of animal in the table is “otter” and the site is “bay”, a monster is drawing a purple rectangle around the row. If those conditions are not met, another monster is putting a line through the column indicating it will be excluded. Stylized text reads “dplyr::filter() - keep rows that satisfy your conditions.”](images/dplyr_filter.jpg){width=100%} 
<p align="center">"Artwork by <a href="https://twitter.com/allison_horst?lang=en">@allison_horst</a>" </p>

If you review the R Documentation for the [filter function](https://www.rdocumentation.org/packages/dplyr/versions/0.7.8/topics/filter) you'll find the following explanation:  

 > Use filter() find rows/cases where conditions are true. Unlike base subsetting with [, rows where the condition evaluates to NA are dropped.  
 
In more plain language, this means that filter() will search through the dataset you provide it, evaluate each row against the provided condition, and then return you a, potentially, smaller dataset with only the rows that met your condition. This is a necessary tool for narrowing down a large dataset into one that you can use. **Keep in mind, this doesn't modify/alter the original dataset!**  


## Practice Dataset  

### Palmer Penquins  

![Palmer Penguins, showing drawings of the Chinstrap (left), Gentoo (center), and Adélie (right) penguins](images/palmer_penguins.png){width=100%} 
<p align="center">"Artwork by <a href="https://twitter.com/allison_horst?lang=en">@allison_horst</a>" </p>

This tutorial's practice exercises will use the [Palmer Penguins Dataset](https://allisonhorst.github.io/palmerpenguins/), which details the recorded attributes of 344 penguins that belong to three different species of penguin, the Chinstrap, the Gentoo, and the Adélie. Following the principles of tidy data, each row represents a single penguin and each column is a variable representing a characteristic of that penguin. If you review the help page (?penguins), you'll find the following list for these variables:  
  - **species** : Chinstrap, Gentoo, or Adélie  
  - **island**: Biscoe, Dream, Torgersen  
  - **bill_length_mm** : The length of the bill (in millimeters) across its top ridge  
  - **bill_depth_mm** : The height of the bill (in millimeters) while closed  
  - **flipper_length_mm** : The length of the flipper, in millimeters  
  - **body_mass_g** : The penguins mass in grams  
  - **sex** : male, female  
  - **year** : The year the penguin data was recorded in (2007, 2008, 2009)  

```{r view_penguins}
head(penguins)
```

Reference: 
Horst AM, Hill AP, Gorman KB (2020). palmerpenguins: Palmer Archipelago (Antarctica) penguin data. R package version 0.1.0. https://allisonhorst.github.io/palmerpenguins/. doi: 10.5281/zenodo.3960218. 

## Filtering on One Condition
<span style="font-weight:bold; font-size:2em"> <span class="underline uOne">subset</span> <- filter(<span class="underline uTwo">data</span>, <span class="underline uThree" >condition</span>)</span>

The standard format for using filter is listed above, which requires you to include 3 things.

<dl>
  <dt class="underline uOne thin">subset</dt>
  <dd> Filter doesn't actually modify the provided dataset, so you'll need to store the results in order to use them in other operations. Later on we'll discuss how we can use the pipe **%>%** operator to remove the need for storing our results! But for now, make sure to always catch your filter results in a variable</dd>
  <dt class="underline uTwo thin">data</dt>
  <dd>Here you provide the name of the dataset you want to work with. For most of our practice exercises, that will be the **penguin** dataset.</dd>
  <dt class="underline uThree thin">condition</dt>
  <dd>This is the most important part for using filter, the condition that determines which rows you want to keep. This expression must evaluate to either true (keep the row) or false (don't keep the row). So always make sure to double check your condition and ensure you aren't missing any rows or accidentally keeping invalid ones. To start with, you'll use one of the basic conditional operators, listed below:
    <table class="table">
      <tr>
        <td colspan="2">**Conditional Statement Syntax: **</td>
        <td colspan="2">**A operand B**</td>
      </tr>
      <tr>
        <td>**==**</td><td> Equal to </td><td>**!=**</td><td> Not Equal to*</td>
      </tr>
      <tr>
        <td>**&lt;**</td><td> Less than </td><td>**&lt;=**</td><td> Less than or Equal to*</td>
      </tr>
      <tr>
        <td>**&gt;**</td><td> Greater than </td><td>**&gt;=**</td><td> Greater than or Equal to*</td>
      </tr>
      <tr>
        <td colspan="2"></td>
        <td colspan="2">**\*Order matters, =>, =<, and =! are not valid syntax!**</td>
      </tr>
    </table>
  </dd>
</dl>   

### Simple Dataset 
```{r, echo=FALSE}
sidebarLayout(
    sidebarPanel(
      selectInput("compareOp_S", "Comparison Operator", c("Equal to (==)", "Not Equal to (!=)", "Less Than (<)", "Less Than or Equal to (<=)", "Greater Than (>)", "Greater Than or Equal to (>=)")),
      div(id="sliderDiv_S", sliderInput("slider_S", "Var Value", value=5, min=0, max=10, step=1))
      ),
  mainPanel(
    htmlOutput("filterOneExpression_S"), 
    htmlOutput("filterOneTable_S")))
```

```{r, context="server"}
expression_S <- reactive(
  switch(input$compareOp_S, 
    "Equal to (==)" = paste("filter(simple, var == <span style=\"color:red\">", input$slider_S, "</span>)", sep=""),
    "Not Equal to (!=)"=paste("filter(penguins, var != <span style=\"color:red\">", input$slider_S, "</span>)", sep=""),
    "Less Than (<)"=paste("filter(simple, var < <span style=\"color:red\">", input$slider_S, "</span>)", sep=""),
    "Less Than or Equal to (<=)"=paste("filter(simple, var <= <span style=\"color:red\">", input$slider_S, "</span>)", sep=""),
    "Greater Than (>)"=paste("filter(simple, var > <span style=\"color:red\">", input$slider_S, "</span>)", sep=""),
    "Greater Than or Equal to (>=)"=paste("filter(simple, var >= <span style=\"color:red\">", input$slider_S, "</span>)", sep=""),
    "None Set"))

output$filterOneExpression_S <- renderUI(HTML(paste("<pre><code>simple <- tibble(var = 0:10)\n",expression_S(),"</code></pre>", sep="")))

# Add CSS Classes to modify the slider so the bar is colored correctly. By default, the number slider fills in on the left slide of 
# the slider's marker, but that doesn't correctly denote all conditions. For example, > should fill in from the right side!
# Current Classes
# noSlider <- removes the blue bar completely (==)
# fullSlider<- completely fills in the blue bar (!=)
# includeSlider <- modifies the marker to denote a value is included (==, <=, >=)
# reverseSlider <- modifies the blue bar to fill in from the right (>, >=)
observeEvent(input$compareOp_S, {
 toggleClass("sliderDiv_S", "fullSlider", condition = input$compareOp_S == "Not Equal to (!=)")
 toggleClass("sliderDiv_S", "noSlider", condition = input$compareOp_S == "Equal to (==)")
 toggleClass("sliderDiv_S", "includeSlider", condition = input$compareOp_S == "Equal to (==)" || input$compareOp_S == "Less Than or Equal to (<=)" || input$compareOp_S == "Greater Than or Equal to (>=)")
 toggleClass("sliderDiv_S", "reverseSlider", condition = input$compareOp_S == "Greater Than (>)" || input$compareOp_S == "Greater Than or Equal to (>=)")
})

#Basic numeric dataset for testing conditional statements
simple <- tibble(var = 0:10)

data_Included <- reactive(
  switch(input$compareOp_S, 
    "Equal to (==)" = filter(simple, var == as.numeric(input$slider_S)),
    "Not Equal to (!=)"= filter(simple, var != as.numeric(input$slider_S)),
    "Less Than (<)"= filter(simple, var < as.numeric(input$slider_S)),
    "Less Than or Equal to (<=)"= filter(simple, var <= as.numeric(input$slider_S)),
    "Greater Than (>)"= filter(simple, var > as.numeric(input$slider_S)),
    "Greater Than or Equal to (>=)"= filter(simple, var >= as.numeric(input$slider_S)),
    simple$var))

data_Excluded<- reactive(
  switch(input$compareOp_S, 
    "Equal to (==)" = filter(simple, var != as.numeric(input$slider_S)),
    "Not Equal to (!=)"= filter(simple, var == as.numeric(input$slider_S)),
    "Less Than (<)"= filter(simple, var >= as.numeric(input$slider_S)),
    "Less Than or Equal to (<=)"= filter(simple, var > as.numeric(input$slider_S)),
    "Greater Than (>)"= filter(simple, var <= as.numeric(input$slider_S)),
    "Greater Than or Equal to (>=)"= filter(simple, var < as.numeric(input$slider_S)),
    c("")))

output$filterOneTable_S <- renderUI(
    HTML(paste("<table class=\"table\"><tr><th>Included Data</th><th style=\"color:red\">Excluded Data</th></tr><tr><td>",
               paste(as.list(data_Included()$var), collapse=", "),
               "</td><td style=\"color:red; font-weight:bold\" >",
              paste(as.list(data_Excluded()$var), collapse=", "),
               "</td></tr></table>", sep="")))

```

### Palmer Penguins Example
<!-- Shiny App for Comparison Operators -->
```{r, echo=FALSE}
sidebarLayout(
    sidebarPanel(
      selectInput("compareOp", "Comparison Operator", c("Equal to (==)", "Not Equal to (!=)", "Less Than (<)", "Less Than or Equal to (<=)", "Greater Than (>)", "Greater Than or Equal to (>=)")),
      conditionalPanel(
        condition = "input.compareOp == \"Equal to (==)\" || input.compareOp == \"Not Equal to (!=)\"",
        selectInput("equalSelect", label = "Penguin Species", choices = levels(penguins$species)),
      ),
      conditionalPanel(
        condition = "input.compareOp == \"Less Than (<)\"",
        sliderInput("lessThanSlider", "Bill Length", value=30, min=0, max=60, step=1),
      ),
      conditionalPanel(
        condition = "input.compareOp == \"Less Than or Equal to (<=)\"",
        div(class="includeSlider", sliderInput("lessEqualSlider", "Flipper Length", value=170, min=170, max=235, step=1)) 
      ),
      conditionalPanel(
        condition = "input.compareOp == \"Greater Than (>)\"",
        div(class="reverseSlider", sliderInput("greaterThanSlider", "Bill Length", value=30, min=0, max=60, step=1)) 
      ),
      conditionalPanel(
        condition = "input.compareOp == \"Greater Than or Equal to (>=)\"",
        div(class="reverseSlider includeSlider", sliderInput("greaterEqualSlider", "Flipper Length", value=170, min=170, max=235, step=1)) 
      )),
  mainPanel(
    htmlOutput("filterOneExpression"), 
    reactableOutput("filterOneTable")))
```

```{r, context="server"}
expression <- reactive(
  switch(input$compareOp, 
    "Equal to (==)" = paste("filter(penguins, species == <span style=\"color:red\">\"", input$equalSelect, "\"</span>)", sep=""),
    "Not Equal to (!=)"=paste("filter(penguins, species != <span style=\"color:red\">\"", input$equalSelect, "\"</span>)", sep=""),
    "Less Than (<)"=paste("filter(penguins, bill_length_mm < <span style=\"color:red\">", input$lessThanSlider, "</span>)", sep=""),
    "Less Than or Equal to (<=)"=paste("filter(penguins, flipper_length_mm <= <span style=\"color:red\">", input$lessEqualSlider, "</span>)", sep=""),
    "Greater Than (>)"=paste("filter(penguins, bill_length_mm > <span style=\"color:red\">", input$greaterThanSlider, "</span>)", sep=""),
    "Greater Than or Equal to (>=)"=paste("filter(penguins, flipper_length_mm >= <span style=\"color:red\">", input$greaterEqualSlider, "</span>)", sep=""),
    "None Set"))

output$filterOneExpression <- renderUI(HTML(paste("<pre><code>",expression(),"</code></pre>", sep="")))

data <- reactive(
  switch(input$compareOp, 
    "Equal to (==)" = filter(penguins, species == input$equalSelect),
    "Not Equal to (!=)"=filter(penguins, species != input$equalSelect),
    "Less Than (<)"=filter(penguins, bill_length_mm < as.numeric(input$lessThanSlider)),
    "Less Than or Equal to (<=)"=filter(penguins, flipper_length_mm <= as.numeric(input$lessEqualSlider)),
    "Greater Than (>)"=filter(penguins, bill_length_mm > as.numeric(input$greaterThanSlider)),
    "Greater Than or Equal to (>=)"=filter(penguins, flipper_length_mm >= as.numeric(input$greaterEqualSlider)),
    penguins))

output$filterOneTable <- renderReactable({
  reactable(
    data(),
    defaultPageSize = 5,
    columns = list(
      species = colDef(minWidth = 100), 
      island = colDef(minWidth = 100),   
      bill_length_mm = colDef(minWidth = 125),
      bill_depth_mm = colDef(minWidth = 125),
      flipper_length_mm = colDef(minWidth = 145),
      body_mass_g = colDef(minWidth = 125),
      sex = colDef(minWidth = 75),
      year= colDef(minWidth = 50)
    ))})

```

### Practice Problems

#### Practice Problem 1
Finish the code block, so that you list only the female penguins
```{r filter_practice1, exercise = TRUE, exercise.blanks = "___+"}
filter(penguins, sex == _____)
```

```{r filter_practice1-solution}
filter(penguins, sex == "female")
```

```{r filter_practice1-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, sex == "female")))
)
```

#### Practice Problem 2
Finish the code block, so that you list only the penguins whose bill length wasn't recorded.
```{r filter_practice2, exercise = TRUE, exercise.blanks = "___+"}
filter(penguins, ____ bill_length_mm ____)
```

```{r filter_practice2-solution}
filter(penguins, is.na(bill_length_mm))
```


```{r filter_practice2-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, is.na(bill_length_mm))))
)
```

#### Practice Problem 3
Write the necessary code so that only penguins with a body mass that is at least 3250 are listed.
```{r filter_practice3, exercise = TRUE}

```

```{r filter_practice3-solution}
filter(penguins, body_mass_g >= 3250)
```


```{r filter_practice3-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, body_mass_g >= 3250)))
)
```

#### Practice Problem 4
Write the necessary code to list all the penguins found on the Torgersen island.
```{r filter_practice4, exercise = TRUE}

```

```{r filter_practice4-solution}
filter(penguins, island == "Torgersen")
```


```{r filter_practice4-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, island == "Torgersen")))
)
```

#### Practice Problem 5
The following code segment returns an empty frame. Fix the code so that it returns the expected dataframe.
```{r filter_practice5, exercise = TRUE}
filter(penguins, species == "adelie")
```

<div id="filter_practice5-hint">
**Hint:** Text searches are case sensitive and the results should have 152 rows.
</div>


```{r filter_practice5-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, species == "Adelie")))
)
```

## Filtering on Multiple Conditions
In the previous section, we learned how to apply filter to a dataset using a single condition but often we will have many conditions to apply. When we need to apply multiple conditions, we turn to the boolean operators! There are three primary boolean operators we will work with, which are AND (&), OR (|), and NOT (!). 

 *Incomplete, Sections to add:*  
   - Math explanation of each operator, i.e. how they evaluate to true/false
   - Common language explanation, e.g. "I want to buy a new BMW that is either a hybrid or an electric" - brand == "BMW" & (type == "hyrbrid" | type == "electric"). Basically look at how we can convert a conversation of needs/requirements into a conditional statement. 
   - Common helper methods - is.na(), %in%
   
  *Fix Broken Widgets*

### Practice Problems

#### Practice Problem 6
Finish the code block, so that you only list the penguins from Torgersen island with a flipper length of exactly 195 mm.
```{r filter_practice6, exercise = TRUE, exercise.blanks = "___+"}
filter(penguins, island ____ ____ flipper_length_mm ____)
```

```{r filter_practice6-solution}
filter(penguins, island == "Torgersen" & flipper_length_mm == 195)
```

```{r filter_practice6-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, island == "Torgersen" & flipper_length_mm == 195)))
)
```

#### Practice Problem 7
Finish the code block, so that you list only the penguins that have a body mass between 3000 and 3500 (inclusive).
```{r filter_practice7, exercise = TRUE, exercise.blanks = "___+"}
filter(penguins, body_mass_g ____ 3000 ____ body_mass_g ____ 3500 )
```

```{r filter_practice7-solution}
filter(penguins, body_mass_g >= 3000 & body_mass_g <= 3500)
```

```{r filter_practice7-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, body_mass_g >= 3000 & body_mass_g <= 3500)))
)
```

#### Practice Problem 8
Write the necessary code to list all of the Gentoo penguins who have measurements entered for both their bill depth and bill length.
```{r filter_practice8, exercise = TRUE}

```

```{r filter_practice8-solution}
filter(penguins, species == "Gentoo", !is.na(bill_depth_mm), !is.na(bill_length_mm))
```

```{r filter_practice8-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, species == "Gentoo", !is.na(bill_depth_mm), !is.na(bill_length_mm))))
)
```

#### Practice Problem 9
Write the necessary code to retrieve all the penguins who had their data recorded in 2007 or 2009, do not live on Biscoe island, and a flipper length between 190 and 200 (exclusive)
```{r filter_practice9, exercise = TRUE}

```

```{r filter_practice9-solution}
filter(penguins, (year == 2007 | year == 2009), island != "Biscoe", flipper_length_mm >190, flipper_length_mm < 200)
```

```{r filter_practice9-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, (year == 2007 | year == 2009), island != "Biscoe", flipper_length_mm >190, flipper_length_mm < 200)))
)
```

#### Practice Problem 10
The following code segment is intended to only list the lighter penguins, but is instead listing all of the penguins, regardless of size. Fix the code, so that only the chosen set of penguins are listed.
```{r filter_practice10, exercise = TRUE}
filter(penguins, body_mass_g < 3500 | body_mass_g > 2000)
```

<div id="filter_practice10-hint">
**Hint:** You should have a result with 71 rows. If you don't, try picking two test values for body mass. One that should pass the condition and one that should fail. Manually test your condition to track down which part might contain the error.
</div>


```{r filter_practice10-check}
grade_result(
  pass_if(~identical(.result, filter(penguins, body_mass_g < 3500 & body_mass_g > 2000)))
)
```

## Homework Problems

#### Problem 1 

```{r filter_hmw1, echo=FALSE}

question("(True or False), filter() modifies the dataframe passed to it.",
  answer("True"),
  answer("False", correct=TRUE),
  incorrect = "Incorrect, as stated [here](https://r4ds.had.co.nz/transform.html?q=filter#filter-rows-with-filter), filter() does not modify the data passed to it. filter() instead creates and returns a new dataframe.",
  allow_retry = TRUE
)
```



#### Problem 2 
```{r filter_hmw2, echo=FALSE, exercise.startover = TRUE}
  question("How many rows will this code segment return? $\\texttt{filter(mpg, displ > 2)}$",
  answer("191", correct = TRUE),
  answer("212"),
  answer("11"),
  answer("0"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

#### Problem 3 
Fill in the blanks with the necessary code so that only the audi cars with a city mpg of 20 or more are listed.

```{r filter_hmw3, exercise = TRUE, exercise.blanks = "___+", exercise.completion = FALSE}
filter(mpg, manufacturer ____ "audi" ____ cty ____ 20)
```

```{r filter_hmw3-check}
grade_result(
  pass_if(~identical(.result, filter(mpg, manufacturer == "audi" & cty >= 20)))
)
```

#### Problem 4 
Debugging Exercise - Fix the following piece of code so that only compact cars are listed in the results.
```{r filter_hmw4, exercise = TRUE, exercise.completion = FALSE}
filter(mpg, CLASS == "COMPACT")
```

```{r filter_hmw4-check}
grade_result(
  pass_if(~identical(.result, filter(mpg, class== "compact")))
)
```

#### Problem 5
Write the necessary code to list all of the cars with an average highway mpg of 20 or more.
```{r filter_hmw5, exercise = TRUE, exercise.completion = FALSE}

```  

```{r filter_hmw5-check}
grade_result(
  pass_if(~identical(.result, filter(mpg, hwy >= 20)))
)
```

#### Problem 6
Write the necessary code to list out all of the minivans and subcompacts produced in 1999.
```{r filter_hmw6, exercise = TRUE, exercise.completion = FALSE}

```

<div id="filter_hmw6-hint">
**Hint:** This problem will require you to use both (&) and (|), so make sure to place them wisely.
</div>

```{r filter_hmw6-check}
grade_result(
  pass_if(~identical(.result, filter(mpg, (class == "minivan" | class == "subcompact") & year == 1999)))
)
```

<!-- Convert to interactive element -->
```{r, echo=FALSE}
fluidRow(
      column(6, align="center", downloadButton("hmw_submission", "Download Homework Submission")),
      column(6, align="center", textOutput("grade")))

```




```{r, context="server"}

  s_id <- reactive({hmw_info$s_id})
  grade <- reactive({hmw_info$hmw_grade})
  homework <- reactive({hmw_info$submissions})
  
  output$grade <- renderText({ 
      paste("Grade: ", format(round((grade() / length(homework())) * 100, 2), nsmall = 2), "%")
    })
  
  
  
  output$hmw_submission <- downloadHandler(
  filename = function() {
    paste0("Filter_Homework_", Sys.Date(), ".html")
  },
  content = function(file) {
    print("Homework Submission: ")
    print(homework())
    rmarkdown::render("../SharedResources/report/hmw_submission.Rmd",
                      output_file = file, 
                      params = list(
                        title = "Filter Homework", 
                        user = s_id(),
                        grade = format(round((grade() / length(homework())) * 100, 2), nsmall = 2),
                        submissions = homework()
                      ),
                      envir = new.env(),
                      intermediates_dir = tempdir())
  }
)

```

